# Intelligent Vehicles (ME41106) / Machine Perception (RO47004) Final Assignment
* (c) TU Delft, 2021
* Period: 2021-2022, Q2

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Moreover, if you see an empty cell, please DO NOT delete it, instead run that cell as you would run all other cells. Please fill in your name(s) and other required details below:

<!-- grade_id: cell-10cbbd705739ac0f, points: 0, grade: false, solution: true -->
```python
STUDENT_NAME = "Job Mulder"
STUDENT_NUMBER = "4538323"
COURSE_TYPE = "MP"  # choose "IV" (ME41106), or "MP" (RO47004)
RUN_ALL = False # You can leave this setting to 'False'
# fill variables above and remove the NotImplementedError() below
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-c35e4c8223095209, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert STUDENT_NAME != ""
assert STUDENT_NUMBER != ""
assert COURSE_TYPE in {"IV", "MP"}, 'Set COURSE_TYPE to "IV" or "MP"'
```

### Pledge
**I will not**
* discuss the content of this assignment with anyone else apart from the instructor team or the Examination Board (upon appeal)
* use solution concepts/code from others. I understand that passing the work of others as my own’s constitutes scientific fraud - even if it is just one line of code (the sole exception to this is the use of concepts/code that were produced in the practica by my own practicum partner; this is allowed). 
* make my own solution concepts/code available to others (neither before nor after the submission deadline) - this includes dissemination via the internet, social media or WhatsApp groups. I understand that unauthorized sharing of solution concepts/code will be considered as abetting scientific fraud.
* re-distribute any parts of the assignment text (neither before nor after the submission deadline) -  this includes dissemination via the internet, social media or WhatsApp groups. I understand that unauthorized re-distribution of the assignment text constitutes copyright infringement. 

**I understand that**
* automatic code plagiarism checks will be employed for all submissions possibly augmented with "manual" checks. I understand that cases with a suspicion of (abetting) scientific fraud will always be forwarded to the Examination Board (penalties include exclusion from the course, inclusion in an  academic dishonesty registry, and even removal from the MSc program)
* cases with copyright infringement will be pursued by the copyright owner(s)

**I agree**
(please sign with your name in cell below)

<!-- grade_id: cell-71566d7fa962d576, points: 0, grade: false, solution: true -->
```python
SIGNED_NAME = "Job Mulder"  # Add your name here to agree to above points (same as STUDENT_NAME above)
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-71566d7fa962d577, points: 0, grade: true, solution: false -->
```python
assert SIGNED_NAME == STUDENT_NAME
```

### General announcements

* There are no groups for the final assignment. It's an _individual_ assessment.
* Academic integrity is a high good at TU Delft. You need to sign the pledge above.
* The assignment is part of the _assessment_, meaning that there is no longer an opportunity to ask questions about the contents on Brightspace or during the lab hours.

### Submission guidelines

* There is a strict deadline for the final assignment (see *Final Assignment* section below). You are responsible to ensure that you have submitted in time. So, please double check that your upload succeeded to Brightspace and that you uploaded the correct files. For late submissions, one grade point is subtracted per day (max. 2 days).

* It is possible to submit multiple versions of the final assignment. However, only your last submission is kept by Brightspace and will be graded. Note that re-submitting after the deadline activates the late-policy!

* **Only submit the `*.ipynb` files and nothing else.** ZIP files are not allowed and will be rejected by Brightspace.
* **Also submit the practica 1, 2 and/or 3 `*.ipynb` files if you are reusing functions from those via `ipynb` imports.**
* You can copy functions/code from your practica. In that case, please **mark the copied code as such** (via a comment). Reuse of code (e.g., via `ipynb` imports) is preferred over code duplication.
* Your code must not use internet connection (to download files or similar).
* Clear the output of your jupyter notebooks via (Cell $\rightarrow$ All Output $\rightarrow$ Clear), and save the notebook afterwards.
* Submit your `release/*/*.ipynb` files in the course you are following. So in case of IV, submit your files on the IV Brightspace page. In case of MP, submit your files on the MP Brightspace page.

### Working with Jupyter

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. You are free to add additional cells, but if you find a cell that you cannot modify or remove, please know that this is on purpose.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for you that you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore. Any grading failure due to such reasons would be considered as a failure of that particular assignment.

* Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All). Also **clear all outputs** from your cells before submitting (select Cell $\rightarrow$ All Output $\rightarrow$ Clear). Please **save the notebook** file afterwards (Ctrl+s, or File $\rightarrow$ Save and Checkpoint)

# Final Assignment


- Deadline: Monday, 10 January 2022 at 10:45 CET
- Dataset: View-of-Delft 
- Sensors: mono camera, LiDAR, radar

## Modules
- [10 pts.] Module 1a "Data Visualization" (**MP and IV**)
- [35 pts.] Module 2a "3D Pedestrian Detection based on a Single Camera" (**MP and IV**)
- [25 pts.] Module 2b "3D Pedestrian Detection based on Multiple Sensors" (**MP only**)
- don't worry about the missing Module 3 :)
- [25 pts.] Module 4 "Motion Planning" (**IV only**)

## Objectives

This assignment makes you apply the knowledge you gained thoughout the three practica to detect pedestrians in 3D.
**IV only**: Given the pedestrian tracks and past ego-vehicle positions, you will plan a safe trajectory for the ego-vehicle.

Compared to the practica, we provide you freedom to choose among concepts that were discussed in the lectures.
This means that we provide you with input data and define the required output format. The rest is up to your creativity and skill.


## Grading
The number of points allotted to the questions more or less corresponds to the time effort involved.

The final grade of this assignment will consist of the following (to some degree) independent criteria:
- Concepts (~30%)
  - Specification of intended solution
  - Proactive reflection
  - Interpretation of experimental results
  - Future work
- Implementation (~40%)
  - Covers the degree to which the earlier proposed concepts are implemented in the given task
- Performance (~20%)
  - Quantitative detection/tracking/motion planning performance (e.g. precision-recall, mAP)
- Efficiency (~10%)
  - Method complexity (efficiency of concept)
  - Runtime (CPU run-time per frame)


For grading, we will run the assignment notebooks in lexicographical order.
We expect each notebook to run from top to bottom without errors.
Please stick to the (de)serialization code we provide.

For any programming task, we recommend to start implementing a basic concept and iteratively refining it, time permitting.
You will be better off with a working implementation of a basic concept rather than with a fancy concept, which ends up non-working.

You should visualize intermediate steps of your implementations to allow the graders to give you partial credit and ease the understanding of your approach.

**Do not spend more than the total time allocated (32 hours) for this final assignment!**
One could spend a massive amount of time to develop truly advanced solutions, but this is neither expected nor will this be favorable in terms of cost-benefit.
There are also other courses to attend to, and having a break in the holiday period is important to regenerate. 

## Some final notes
Please complete the assignment notebooks in the following lexicographic order, i.e.,
- **MP and IV**  `fa_01a_data_visualization.ipynb`
- **MP and IV**: `fa_02a_3d_pedestrian_detection_single_camera.ipynb`
- **MP only**:   `fa_02b_mp_only_3d_pedestrian_detection_multiple_sensors.ipynb`
- again, don't worry about the missing 03 :)
- **IV only**:   `fa_04_iv_only_motion_planning.ipynb`




### Code reuse
You are encouraged to re-use code you implemented **in your own group** within the practica.
E.g., you can import the function `project_points` from Practicum 1 by the statement `from ipynb.fs.defs.practicum1 import project_points`.
It can be fixed implementations after you discussed the solutions within the practicum sessions.
If you run into any errors when importing via `ipynb`, please make sure that you don't have variable names in the to-be-imported notebooks which consist of pure capital letters (such as `T`, or `XYZ`, but also `P2`).
See [ipynb docs](https://ipynb.readthedocs.io/en/stable/#import-only-definitions).
Sadly, the errors thrown by `ipynb` are not very helpful, so go check for syntactic errors and capital letter variables.
Try restarting the kernel via `Kernel` -> `Restart and Clear Output` and run the code again starting from the topmost cell.
You can also copy over code from the jupyter notebooks of your practica to the final assignment, but please mark them as such.

You can use any python modules which are available in the conda `environment.yml` (and dependencies thereof).
You are not allowed to use additional python modules/packages or have code modified outside the jupyter notebooks.
If you think a change outside the jupyter notebooks makes sense, please get in contact with us (see below).

### Questions
Since this final assignment is part of the assessment, we will no longer answer methodical questions.
If you have any other issues (unclarity in the phrasing of the task, bugs, data problems, infrastructural issues) please [get in contact with Ronald Ensing by email](mailto:R.M.Ensing@tudelft.nl).
Do not use a Brightspace forum for the Final Assignment!

```python

```